/*heap*/
// I AM NOT DONE

use std::cmp::Ord;
use std::default::Default;
pub struct Heap<T>
where
    T: Default,
{
    count: usize,
    v: Vec<T>,
    comparator: fn(&T, &T) -> bool,
}
impl<T> Heap<T>
where
    T: Default,
{
    pub fn new(comparator: fn(&T, &T) -> bool, v:Vec<T>) -> Self {
        Self {
            count: 0,
            v: v,
            comparator,
        }
    }
    pub fn len(&self) -> usize {
        self.count
    }
    pub fn is_empty(&self) -> bool {
        self.len() == 0
    }
    pub fn add(&mut self, value: T) {
        self.count += 1;
        self.v.push(value);
        // Heapify Up
        let mut idx = self.count;
        while self.parent_idx(idx) > 0 {
            let pdx = self.parent_idx(idx);
            if (self.comparator)(&self.v[idx], &self.v[pdx]) {
                self.v.swap(idx, pdx);
            }
            idx = pdx;
        }
    }
	pub fn get_vec(&self) -> Vec<T>{
		self.v
	}
    fn parent_idx(&self, idx: usize) -> usize {
        idx / 2
    }
    fn children_present(&self, idx: usize) -> bool {
        self.left_child_idx(idx) <= self.count
    }
    fn left_child_idx(&self, idx: usize) -> usize {
        idx * 2
    }
    fn right_child_idx(&self, idx: usize) -> usize {
        self.left_child_idx(idx) + 1
    }
    fn smallest_child_idx(&self, idx: usize) -> usize {
        if self.right_child_idx(idx) > self.count {
            self.left_child_idx(idx)
        } else {
            let ldx = self.left_child_idx(idx);
            let rdx = self.right_child_idx(idx);
            if (self.comparator)(&self.v[ldx], &self.v[rdx]) {
                ldx
            } else {
                rdx
            }
        }
    }
}
impl<T> Heap<T>
where
    T: Default + Ord,
{
    /// Create a new MinHeap
    pub fn new_min() -> Self {
        Self::new(|a, b| a < b, vec![])
    }
    /// Create a new MaxHeap
    pub fn new_max() -> Self {
        Self::new(|a, b| a > b, vec![])
    }
}
impl<T> Iterator for Heap<T>
where
    T: Default,
{
    type Item = T;
    fn next(&mut self) -> Option<T> {
        if self.count == 0 {
            return None;
        }
        // This feels like a function built for heap impl :)
        // Removes an item at an index and fills in with the last item
        // of the Vec
        let next = Some(self.v.swap_remove(1));
        self.count -= 1;
        if self.count > 0 {
            // Heapify Down
            let mut idx = 1;
            while self.children_present(idx) {
                let cdx = self.smallest_child_idx(idx);
                if !(self.comparator)(&self.v[idx], &self.v[cdx]) {
                    self.v.swap(idx, cdx);
                }
                idx = cdx;
            }
        }
        next
    }
}
pub struct MinHeap;
impl MinHeap {
    #[allow(clippy::new_ret_no_self)]
    pub fn new<T>(v:Vec<T>) -> Heap<T>
    where
        T: Default + Ord,
    {
        Heap::new(|a, b| a < b,v)
    }
}
pub struct MaxHeap;
impl MaxHeap {
    #[allow(clippy::new_ret_no_self)]
    pub fn new<T>(v:Vec<T>) -> Heap<T>
    where
        T: Default + Ord,
    {
        Heap::new(|a, b| a > b,v)
    }
}
#[cfg(test)]
mod tests {
    use super::*;
    #[test]
    fn heap_test1() {
        let mut heap = MaxHeap::new::<i32>(vec![]);
        assert_eq!(heap.next(), None);
    }
    #[test]
    fn heap_test2() {
        let mut heap = MinHeap::new(vec![]);
        assert_eq!(heap.next(), Some(1));
    }
    #[test]
    fn heap_test3() {
        let mut heap = MaxHeap::new(vec![]);
        heap.add(4);
        heap.add(2);
        heap.add(9);
        heap.add(11);
        assert_eq!(heap.len(), 4);
        assert_eq!(heap.next(), Some(11));
        assert_eq!(heap.next(), Some(9));
        assert_eq!(heap.next(), Some(4));
        heap.add(1);
        assert_eq!(heap.next(), Some(2));
    }
}